# Accidentes de tráfico en Reino Unido entre 2010 y 2014 

### Disponible en Kaggle en:
https://www.kaggle.com/stefanoleone992/adm-project-road-accidents-in-uk

### Variables y significado

* Accident_Index: Accident index
* Latitude: Accident latitude
* Longitude: Accident longitude
* Region: Accident region
* Urban_or_Rural_Area: Accident area (rural or urban)
* X1st_Road_Class: Accident road class
* Driver_IMD_Decile: Road IMD Decile
* Speed_limit: Road speed limit
* Road_Type: Road type
* Road_Surface_Conditions: Road surface condition
* Weather: Weather
* High_Wind: High wind
* Lights: Road lights
* Datetime: Accident datetime
* Year: Accident year
* Season: Accident season
* Month_of_Year: Accident month
* Day_of_Month: Accident day of month
* Day_of_Week: Accident day of week
* Hour_of_Day: Accident hour of day
* Number_of_Vehicles: Accident number of vehicles
* Age_of_Driver: Driver age
* Age_of_Vehicle: Vehicle age
* Junction_Detail: Accident junction detail
* Junction_Location: Accident junction location
* X1st_Point_of_Impact: Vehicle first point of impact
* Driver_Journey_Purpose: Driver journey purpose
* Engine_CC: Vehicle engine power (in CC)
* Propulsion_Code: Vehicle propulsion code
* Vehicle_Make: Vehicle brand
* Vehicle_Category: Vehicle brand category
* Vehicle_Manoeuvre: Vehicle manoeuvre when accident happened
* Accident_Severity: Accident severity

**Nombre completo del alumno:**  CRISTINA RODRIGUEZ CHAMORRO - DNI: 77380999Y

**INSTRUCCIONES**: en cada celda debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) que por defecto vienen inicializadas a `None`. No se necesita usar variables intermedias, pero puedes hacerlo siempre que el resultado final del cálculo quede guardado exactamente en la variable que venía inicializada a None (debes reemplazar None por la secuencia de transformaciones necesarias, pero nunca cambiar el nombre de esa variable). **No olvides borrar la línea *raise NotImplementedError()* de cada celda cuando hayas completado la solución de esa celda y quieras probarla**.

Después de cada celda evaluable verás una celda con código. Ejecútala (no modifiques su código) y te dirá si tu solución es correcta o no. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero evaluar dicha celda es un indicador bastante fiable acerca de si realmente has implementado la solución correcta o no. Asegúrate de que, al menos, todas las celdas indican que el código es correcto antes de enviar el notebook terminado.

### Sobre el dataset anterior (accidents_uk.csv) se pide:

* Leerlo tratando de que Spark infiera el tipo de dato de cada columna, y cachearlo.

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR LAS VARIABLES
accidentesDF = spark.read \
                 .option("header", "true") \
                 .option("inferSchema", "true") \
                 .csv("gs://ntic-bucket-cristina/accidents_uk.csv").cache()
accidentesDF.show()
```

Output:
```
+--------------+---------+---------+------+-------------------+---------------+-----------------+-----------+------------------+-----------------------+-------+---------+-----------------+-------------------+----+------+-------------+------------+-----------+-----------+------------------+-------------+--------------+--------------------+--------------------+--------------------+----------------------+---------+---------------+------------+----------------+-------------------+-----------------+
|Accident_Index| Latitude|Longitude|Region|Urban_or_Rural_Area|X1st_Road_Class|Driver_IMD_Decile|Speed_limit|         Road_Type|Road_Surface_Conditions|Weather|High_Wind|           Lights|           Datetime|Year|Season|Month_of_Year|Day_of_Month|Day_of_Week|Hour_of_Day|Number_of_Vehicles|Age_of_Driver|Age_of_Vehicle|     Junction_Detail|   Junction_Location|X1st_Point_of_Impact|Driver_Journey_Purpose|Engine_CC|Propulsion_Code|Vehicle_Make|Vehicle_Category|  Vehicle_Manoeuvre|Accident_Severity|

```

<!-- grade_id: read_csv_test, points: 1, grade: true, solution: false -->
```python
from pyspark.sql.types import DoubleType
assert(accidentesDF.schema[1].dataType == DoubleType())
```

* Discretizar la variable **Age_of_Vehicle** utilizando un bucketizer (sin crear un pipeline) en los puntos de corte (0, 2, 5, 10, 15, 20, 35). La discretización debe quedar en una nueva columna de tipo Double llamada **Edad_Vehiculo**.

<!-- grade_id: bucketize, points: 0, grade: false, solution: true -->
```python
from pyspark.ml.feature import Bucketizer
# LÍNEAS EVALUABLES, NO RENOMBRAR LAS VARIABLES
splits = [0.0,2.0,5.0,10.0,15.0,20.0,35.0, float("inf")]
bucketizer = Bucketizer(splits = splits, inputCol = "Age_of_Vehicle",outputCol = "Edad_Vehiculo")
accidentesBucketizedDF = bucketizer.transform(accidentesDF)
accidentesBucketizedDF.show()
```

Output:
```
+--------------+---------+---------+------+-------------------+---------------+-----------------+-----------+------------------+-----------------------+-------+---------+-----------------+-------------------+----+------+-------------+------------+-----------+-----------+------------------+-------------+--------------+--------------------+--------------------+--------------------+----------------------+---------+---------------+------------+----------------+-------------------+-----------------+-------------+
|Accident_Index| Latitude|Longitude|Region|Urban_or_Rural_Area|X1st_Road_Class|Driver_IMD_Decile|Speed_limit|         Road_Type|Road_Surface_Conditions|Weather|High_Wind|           Lights|           Datetime|Year|Season|Month_of_Year|Day_of_Month|Day_of_Week|Hour_of_Day|Number_of_Vehicles|Age_of_Driver|Age_of_Vehicle|     Junction_Detail|   Junction_Location|X1st_Point_of_Impact|Driver_Journey_Purpose|Engine_CC|Propulsion_Code|Vehicle_Make|Vehicle_Category|  Vehicle_Manoeuvre|Accid
```

<!-- grade_id: bucketize_tests, points: 1, grade: true, solution: false -->
```python
assert("Edad_Vehiculo" in accidentesBucketizedDF.columns)
assert(accidentesBucketizedDF.schema.fields)
```

* Crear un nuevo DF donde la columna "Age_of_Driver" ha sido reemplazada por otra de tipo string en la que los valores 1 y 2 son "Adolescente", los valores 3 y 4 por "Joven", los valores 5 y 6 por "Adulto", y los demás valores se dejan sin modificar.

<!-- grade_id: renombrar_edad, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES

from pyspark.sql.functions import col, StringType, when
accidentesAge = accidentesDF.withColumn("Age_of_Driver", col("Age_of_Driver").cast(StringType()))
accidentesAgeDF = accidentesAge.withColumn("Age_of_Driver", when((col("Age_of_Driver") == "1") | (col("Age_of_Driver") == "2"), "Adolescente")\
                                         .when((col("Age_of_Driver") == "3") | (col("Age_of_Driver") == "4"), "Joven")\
                                         .when((col("Age_of_Driver") == "5") | (col("Age_of_Driver") == "6"), "Adulto")\
                                          .when((col("Age_of_Driver") == "7"), "7")\
                                          .when((col("Age_of_Driver") == "8"), "8"))
accidentesAgeDF.show()
```

Output:
```
+--------------+---------+---------+------+-------------------+---------------+-----------------+-----------+------------------+-----------------------+-------+---------+-----------------+-------------------+----+------+-------------+------------+-----------+-----------+------------------+-------------+--------------+--------------------+--------------------+--------------------+----------------------+---------+---------------+------------+----------------+-------------------+-----------------+
|Accident_Index| Latitude|Longitude|Region|Urban_or_Rural_Area|X1st_Road_Class|Driver_IMD_Decile|Speed_limit|         Road_Type|Road_Surface_Conditions|Weather|High_Wind|           Lights|           Datetime|Year|Season|Month_of_Year|Day_of_Month|Day_of_Week|Hour_of_Day|Number_of_Vehicles|Age_of_Driver|Age_of_Vehicle|     Junction_Detail|   Junction_Location|X1st_Point_of_Impact|Driver_Journey_Purpose|Engine_CC|Propulsion_Code|Vehicle_Make|Vehicle_Category|  Vehicle_Manoeuvre|Accident_Severity|

```

<!-- grade_id: renombrar_edad_test, points: 1, grade: true, solution: false -->
```python
assert(dict(accidentesAgeDF.dtypes)["Age_of_Driver"] == "string")
collectedDF = accidentesAgeDF.groupBy("Age_of_Driver").count().orderBy("count").collect()
print(collectedDF)
assert((collectedDF[0]["count"] == 9195) & (collectedDF[0]["Age_of_Driver"] == "8"))
assert((collectedDF[1]["count"] == 13338) & (collectedDF[1]["Age_of_Driver"] == "7"))
assert((collectedDF[2]["count"] == 57174) & (collectedDF[2]["Age_of_Driver"] == "Adolescente"))
assert((collectedDF[3]["count"] == 67138) & (collectedDF[3]["Age_of_Driver"] == "Adulto"))
assert((collectedDF[4]["count"] == 104987) & (collectedDF[4]["Age_of_Driver"] == "Joven"))
```

Output:
```
[Row(Age_of_Driver=u'8', count=9195), Row(Age_of_Driver=u'7', count=13338), Row(Age_of_Driver=u'Adolescente', count=57174), Row(Age_of_Driver=u'Adulto', count=67138), Row(Age_of_Driver=u'Joven', count=104987)]
```

* Partiendo de `accidentesDF`, crear un nuevo DataFrame de una sola fila que contenga, **por este orden de columnas**, el **número** de categorías existentes para el propósito del viaje, para el tipo de maniobra del vehículo, para las condiciones de la calzada y para la severidad del accidente. Pista: crear las columnas al vuelo con `select`(). Renombrar cada columna de conteo para que se llame igual que la propia columna que estamos contando.

<!-- grade_id: numero_categorias, points: 0, grade: false, solution: true -->
```python
from pyspark.sql.functions import countDistinct
numeroCategoriasDF = accidentesDF\
.select(col("Driver_Journey_Purpose"), col("Vehicle_Manoeuvre"),col("Road_Surface_Conditions"), col("Accident_Severity"))\
.agg(countDistinct(col("Driver_Journey_Purpose"))\
     .alias("Driver_Journey_Purpose"),countDistinct(col("Vehicle_Manoeuvre"))\
     .alias("Vehicle_Manoeuvre"),countDistinct(col("Road_Surface_Conditions"))\
     .alias("Road_Surface_Conditions"),countDistinct(col("Accident_Severity"))\
     .alias("Accident_Severity"))
numeroCategoriasDF.show()
```

Output:
```
+----------------------+-----------------+-----------------------+-----------------+
|Driver_Journey_Purpose|Vehicle_Manoeuvre|Road_Surface_Conditions|Accident_Severity|
+----------------------+-----------------+-----------------------+-----------------+
|                     5|               11|                      5|                2|
+----------------------+-----------------+-----------------------+-----------------+
```

<!-- grade_id: numero_categorias_test, points: 1, grade: true, solution: false -->
```python
assert(len(numeroCategoriasDF.columns) == 4)
assert(numeroCategoriasDF.count() == 1)
```

* Partiendo de `accidentesAgeDF` definido anteriormente, crear un nuevo DataFrame con tantas filas como posibles propósitos de un viaje, y tantas columnas como rangos de edad habíamos distinguido en dicho DataFrame. Las columnas deben llamarse igual que cada uno de los niveles posibles de rangos de edad. Cada casilla del nuevo DataFrame deberá contener el **porcentaje** del número de accidentes ocurridos en ese tipo de viaje (fila) para ese rango de edad (columna), medido sobre el *total de accidentes ocurridos para ese tipo de viaje*.

Pista: se puede hacer todo en una sola secuencia de transformaciones sin variable auxiliar. Calcular primero el conteo, después añadir una columna con los totales de cada tipo de viaje como la suma de las 5 columnas de conteos, y finalmente reemplazar cada columna de conteo por su porcentaje. No debe utilizarse `when` en ningún momento, solo aritmética de columnas. Recuerda cómo desplegar grupos en varias columnas.

<!-- grade_id: viajes_por_edad, points: 0, grade: false, solution: true -->
```python
from pyspark.sql.functions import count

viajesPorEdadDF = accidentesAgeDF\
.groupBy("Driver_Journey_Purpose")\
.pivot("Age_of_Driver")\
.agg(count(col("Driver_Journey_Purpose")))\
.withColumn("Sumas", col("Adolescente") + col("Adulto") + col("Joven") + col("7") + col("8"))\
.withColumn("Adolescente", col("Adolescente")/col("Sumas"))\
.withColumn("Adulto", col("Adulto")/col("Sumas"))\
.withColumn("Joven", col("Joven")/col("Sumas"))\
.withColumn("7", col("7")/col("Sumas"))\
.withColumn("8", col("8")/col("Sumas"))\
.drop(col("Sumas"))
viajesPorEdadDF.show()
```

Output:
```
+----------------------+--------------------+--------------------+-------------------+-------------------+-------------------+
|Driver_Journey_Purpose|                   7|                   8|        Adolescente|             Adulto|              Joven|
+----------------------+--------------------+--------------------+-------------------+-------------------+-------------------+
|  Taking pupil to/f...|0.026578073089700997|0.009060706735125339|0.07550588945937783|0.22259136212624583|   0.66626396858955|
|  Pupil riding to/f...|0.021611001964636542| 0.01768172888015717| 0.6227897838899804|0.12573673870333987|0.21218074656188604|
|  Journey as part o...|0.023330770595540836|0.003814536637293749|0.15536105032822758|0.31802590336507186|  0.499467739073866|
|       Other/Not known| 0.06503169065192967| 0.04814613718101034|0.24026836256901252| 0.2565088137105748|0.39004499588747266|
|  Commuting to/from...|0.012527948326649396|0.002519785640770...|  0.236327501153423| 0.2791993469851297|0.469
```

<!-- grade_id: viajes_por_edad_test, points: 3, grade: true, solution: false -->
```python
assert(len(viajesPorEdadDF.columns) >= 6)
assert("7" in viajesPorEdadDF.columns)
assert("8" in viajesPorEdadDF.columns)
assert("Adolescente" in viajesPorEdadDF.columns)
assert("Joven" in viajesPorEdadDF.columns)
assert("Adulto" in viajesPorEdadDF.columns)
assert(viajesPorEdadDF.columns[0] == "Driver_Journey_Purpose")
assert(viajesPorEdadDF.count() == 5)
commuting = viajesPorEdadDF.orderBy("Driver_Journey_Purpose").collect()[0]
assert(commuting.Driver_Journey_Purpose.startswith("Commuting"))
assert(abs(commuting['7'] - 0.012527948326649396) < 0.001)
assert(abs(commuting['8'] - 0.002519785640770) < 0.001)
assert(abs(commuting.Adolescente - 0.236327501153423) < 0.001)
assert(abs(commuting.Adulto - 0.2791993469851297) < 0.001)
assert(abs(commuting.Joven - 0.46942541789402703) < 0.001)
```

* Unir la información obtenida en el paso anterior al DataFrame `accidentesAgeDF`, de manera que **al resultado final se añada una columna nueva llamada `Porcentaje`** que contenga el porcentaje de accidentes que ha habido para ese rango de edad y ese tipo de viaje de entre todos los viajes que ha habido de ese tipo (es decir, el porcentaje adecuado de la tabla anterior). Por ejemplo, si el accidente se produjo en un trayecto de tipo `Commuting...` y la persona es `Joven`, entonces la columna Porcentaje tomará el valor de la columna `Joven` y por tanto tendrá el valor 0.46942, pero si la persona es `Adulto`, entonces tomará el valor de la columna `Adulto` el cual será 0.2791993469851297.

PISTA: unir los dos DF mediante join() convencional, y a continuación, crear la nueva columna `Porcentaje` en el resultado, utilizando `when` para ver cuál es el valor que debe tener en cada fila (más concretamente: de qué columna debemos tomarlo) en función del valor de la columna `Age_of_Driver`. No se necesitan variables intermedias; se puede hacer en una secuencia de transformaciones encadenadas.

<!-- grade_id: join, points: 0, grade: false, solution: true -->
```python
from pyspark.sql import functions as F
finalDF = viajesPorEdadDF\
.join(accidentesAgeDF, 
      accidentesAgeDF.Driver_Journey_Purpose == viajesPorEdadDF.Driver_Journey_Purpose)\
.withColumn("Porcentaje", when(col("Age_of_Driver") == "7" , F.col("7"))\
           .when(col("Age_of_Driver") == "8", F.col("8"))\
           .when(col("Age_of_Driver") == "Adolescente", F.col("Adolescente"))\
           .when(col("Age_of_Driver") == "Adulto", F.col("Adulto"))\
           .when(col("Age_of_Driver") == "Joven", F.col("Joven")))
finalDF.show()
```

Output:
```
+----------------------+--------------------+--------------------+-------------------+-------------------+----------------+--------------+---------+---------+------+-------------------+---------------+-----------------+-----------+------------------+-----------------------+-------+---------+--------------------+-------------------+----+------+-------------+------------+-----------+-----------+------------------+-------------+--------------+--------------------+--------------------+--------------------+----------------------+---------+---------------+------------+----------------+-------------------+-----------------+--------------------+
|Driver_Journey_Purpose|                   7|                   8|        Adolescente|             Adulto|           Joven|Accident_Index| Latitude|Longitude|Region|Urban_or_Rural_Area|X1st_Road_Class|Driver_IMD_Decile|Speed_limit|         Road_Type|Road_Surface_Conditions|Weather|High_Wind|              Lights|           Datetime|Year|Season|Month_of_
```

<!-- grade_id: join_tests, points: 3, grade: true, solution: false -->
```python
def sum_cond(df, column, condition): 
    return(df.where(condition).select(F.sum(column).alias(column)).collect()[0][column])
    
assert("Porcentaje" in finalDF.columns)
assert(abs(sum_cond(finalDF, "Porcentaje", F.col("Age_of_Driver") == "Adolescente") - 13344.826819125037) < 0.001)
assert(abs(sum_cond(finalDF, "Porcentaje", F.col("Age_of_Driver") == "Joven") - 44438.00809518224) < 0.001)
assert(abs(sum_cond(finalDF, "Porcentaje", F.col("Age_of_Driver") == "Adulto") - 18028.24488479408) < 0.001)
assert(abs(sum_cond(finalDF, "Porcentaje", F.col("Age_of_Driver") == "7") - 812.0952970292334) < 0.001)
assert(abs(sum_cond(finalDF, "Porcentaje", F.col("Age_of_Driver") == "8") - 432.2987413617681) < 0.001)
```